# Loops

Topics:
1. `while` loops
2. `for` loops
<br>

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```

For example, we could use `while` to count or to iterate over an array.

```julia
n = 0
while n < 10
    n += 1
    println(n)
end

```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

i = 1
while i <= length(myfriends)
    friend = myfriends[i]
    println("Hi $friend, it's great to see you!")
    i += 1
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

We could use a for loop to generate the same results as either of the examples above:

```julia
for n in 1:10
    println(n)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

for friend in myfriends
    println("Hi $friend, it's great to see you!")
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

Now let's use `for` loops to create some addition tables, where the value of every entry is the sum of its row and column indices. <br>

Note that we iterate over this array via column-major loops in order to get the best performance. More information about fast indexing of multidimensional arrays inside nested loops can be found at https://docs.julialang.org/en/v1/manual/performance-tips/#Access-arrays-in-memory-order,-along-columns-1

First, we initialize an array with zeros.

```julia
m, n = 100, 100

```

Output:
```
(100, 100)
```

```julia
A = fill(0, (m, n))
```

Output:
```
100×100 Matrix{Int64}:
 0  0  0  0  0  0  0  0  0  0  0  0  0  …  0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0  …  0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0  …  0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0   
```

```julia
for j in 1:n
    for i in 1:m
        A[i, j] = i + j
    end
end
A
```

Output:
```
100×100 Matrix{Int64}:
   2    3    4    5    6    7    8    9  …   95   96   97   98   99  100  101
   3    4    5    6    7    8    9   10      96   97   98   99  100  101  102
   4    5    6    7    8    9   10   11      97   98   99  100  101  102  103
   5    6    7    8    9   10   11   12      98   99  100  101  102  103  104
   6    7    8    9   10   11   12   13      99  100  101  102  103  104  105
   7    8    9   10   11   12   13   14  …  100  101  102  103  104  105  106
   8    9   10   11   12   13   14   15     101  102  103  104  105  106  107
   9   10   11   12   13   14   15   16     102  103  104  105  106  107  108
  10   11   12   13   14   15   16   17     103  104  105  106  107  108  109
  11   12   13   14   15   16   17   18     104  105  106  107  108  109  110
  12   13   14   15   16   17   18   19  …  105  106  107  108  109  110  111
  13   14   15   16   17   18   19   20     106  107  108  109  110  111  112
  14   15   16   17   18   19   20   21  
```

Here's some syntactic sugar for the same nested `for` loop

```julia
B = fill(0, (m, n))
```

Output:
```
100×100 Matrix{Int64}:
 0  0  0  0  0  0  0  0  0  0  0  0  0  …  0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0  …  0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0  …  0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0     0  0  0  0  0  0  0  0  0  0  0  0
 0  0  0  0  0  0  0  0  0  0  0  0  0   
```

```julia
for j in 1:n, i in 1:m
    B[i, j] = i + j
end
B
```

Output:
```
100×100 Matrix{Int64}:
   2    3    4    5    6    7    8    9  …   95   96   97   98   99  100  101
   3    4    5    6    7    8    9   10      96   97   98   99  100  101  102
   4    5    6    7    8    9   10   11      97   98   99  100  101  102  103
   5    6    7    8    9   10   11   12      98   99  100  101  102  103  104
   6    7    8    9   10   11   12   13      99  100  101  102  103  104  105
   7    8    9   10   11   12   13   14  …  100  101  102  103  104  105  106
   8    9   10   11   12   13   14   15     101  102  103  104  105  106  107
   9   10   11   12   13   14   15   16     102  103  104  105  106  107  108
  10   11   12   13   14   15   16   17     103  104  105  106  107  108  109
  11   12   13   14   15   16   17   18     104  105  106  107  108  109  110
  12   13   14   15   16   17   18   19  …  105  106  107  108  109  110  111
  13   14   15   16   17   18   19   20     106  107  108  109  110  111  112
  14   15   16   17   18   19   20   21  
```

The more "Julia" way to create this addition table would have been with an *array comprehension*.

```julia
C = [i + j for i in 1:m, j in 1:n]
```

Output:
```
100×100 Matrix{Int64}:
   2    3    4    5    6    7    8    9  …   95   96   97   98   99  100  101
   3    4    5    6    7    8    9   10      96   97   98   99  100  101  102
   4    5    6    7    8    9   10   11      97   98   99  100  101  102  103
   5    6    7    8    9   10   11   12      98   99  100  101  102  103  104
   6    7    8    9   10   11   12   13      99  100  101  102  103  104  105
   7    8    9   10   11   12   13   14  …  100  101  102  103  104  105  106
   8    9   10   11   12   13   14   15     101  102  103  104  105  106  107
   9   10   11   12   13   14   15   16     102  103  104  105  106  107  108
  10   11   12   13   14   15   16   17     103  104  105  106  107  108  109
  11   12   13   14   15   16   17   18     104  105  106  107  108  109  110
  12   13   14   15   16   17   18   19  …  105  106  107  108  109  110  111
  13   14   15   16   17   18   19   20     106  107  108  109  110  111  112
  14   15   16   17   18   19   20   21  
```

### Exercises

#### 4.1 
Loop over integers between 1 and 100 and print their squares.

```julia
num = 100
squares = [i^2 for i in 1:num]
```

Output:
```
100-element Vector{Int64}:
     1
     4
     9
    16
    25
    36
    49
    64
    81
   100
   121
   144
   169
     ⋮
  7921
  8100
  8281
  8464
  8649
  8836
  9025
  9216
  9409
  9604
  9801
 10000
```

```julia
square_s = fill(0, num)
for i in 1:num
    square_s[i] = i^2
end
square_s
```

Output:
```
100-element Vector{Int64}:
     1
     4
     9
    16
    25
    36
    49
    64
    81
   100
   121
   144
   169
     ⋮
  7921
  8100
  8281
  8464
  8649
  8836
  9025
  9216
  9409
  9604
  9801
 10000
```

#### 4.2 
Add to the code above a bit to create a dictionary, `squares` that holds integers and their squares as key, value pairs such that

```julia
squares[10] == 100
```

```julia
squares = Dict(i => i^2 for i in 1:num)
```

Output:
```
Dict{Int64, Int64} with 100 entries:
  5  => 25
  56 => 3136
  35 => 1225
  55 => 3025
  60 => 3600
  30 => 900
  32 => 1024
  6  => 36
  67 => 4489
  45 => 2025
  73 => 5329
  64 => 4096
  90 => 8100
  4  => 16
  13 => 169
  54 => 2916
  63 => 3969
  86 => 7396
  91 => 8281
  62 => 3844
  58 => 3364
  52 => 2704
  12 => 144
  28 => 784
  75 => 5625
  ⋮  => ⋮
```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia
squares_arr = [i^2 for i in 1:num]
```

Output:
```
100-element Vector{Int64}:
     1
     4
     9
    16
    25
    36
    49
    64
    81
   100
   121
   144
   169
     ⋮
  7921
  8100
  8281
  8464
  8649
  8836
  9025
  9216
  9409
  9604
  9801
 10000
```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```